## ECE 161C Laboratory Notebooks
---

This Juypter notebook is part of a series of computer laboratories which are designed
to teach image-processing programming. They are the practical complement of the theoretical lectures of the UCSD's course 
[**ECE 161C**](https://canvas.ucsd.edu/courses/62410) taught by Prof. Rahul Parhi.

---
# Lab 0: Introduction
**Released**: Wednesday, April 1, 2026
**Due**: Monday, April 6, 2026 at 11:55pm

# Introduction 

This introductory lab will guide you through the tools that you will need for the labs. 

## Table of contents
1. [Introduction to Jupyter notebooks](#1.-Introduction-to-Jupyter-notebooks) 
2. [Elements of Python](#2.-Elements-of-Python)
    1. [Basic Python programming](#2.1.-Basic-Python-programming)
    2. [NumPy](#2.2.-NumPy)
    3. [Matplotlib](#2.3.-Matplotlib)
3. [Image Thresholding](#3.-Image-Thresholding)
    1. [Range Thresholding](#3.1-Range-thresholding)
4. [Image processing in Python using scikit-image](#4.-Image-processing-in-Python-using-scikit-image)
    1. [Examples](#4.1.-Examples)
    2. [Try it yourself!](#4.2.-Try-it-yourself!) 


# 1. Introduction to Jupyter notebooks

[Back to table of contents](#Table-of-contents)

A Jupyter notebook consists of cells that contain either code snippets with their results or text. The cellular structure of Jupyter notebooks allows us to run small snippets of code one after the other, go back and forth between, or check intermediate results between them. 

<b>Note</b>: If you declare some variable in a cell, this variable will be available in all other cells that use the same programming language, irrespective of the location of the cell. Let's look at an example. Read the description of the cells below and run them as instructed.

In this first cell we declare a variable `a` and assign it the value $5$. Run it.

```python
# cell 1
# this is our variable
a = 5
```

In the second cell we print the value of `a`. If you run it just after the cell above, it should print "The value of a is 5". Run it.

```python
# cell 2
# here we print the value of a
print(f"The value of a is {a}")
```

Then, in the third cell, we modify `a` by adding $2$ to it. Run it.

```python
# cell 3
# add two to a
a = a + 2
```

Now run _cell 2_ again: the value of `a` has been modified, even though _cell 2_ is located above _cell 3_ in our notebook. Now run _cell 1_ again, where we set `a=5`, and `a` will be $5$ in all cells again. 
This shows the **global** impact every single cell has on the notebook. 

<div class="alert alert-info">

**Note:** If at any point the kernel should crash for whatever reason, you can restart it by clicking on the button `Kernel` in the toolbar on top and then click on `Restart Kernel...`. Once the kernel is restarted, you will need to run all necessary cells again to regain the working environment you had before. A good practice to test that there are no errors in the notebook is to click `Restart and Run All Cells...` once you have finished a notebook. This will run all the cells from top to bottom and only stop if an error occurs.
</div>

# 2. Elements of Python

[Back to table of contents](#Table-of-contents)

_Python_ is a high-level language with many great libraries with built-in image processing algorithms. You will learn how to apply the different image processing methods to images using the fastest and most standard tools and visualize the results. 

    
**Note:** If you are already familiar with [Python](python.org) and the [NumPy](numpy.org) and [Matplotlib](matplotlib.org) libraries, you can jump to the next section.
</div>

## 2.1. Basic Python programming
[Back to table of contents](#Table-of-contents)

There are three basic types of variables  in Python:
```python
# Numbers
my_var = 2
my_var = 5.2
# Strings
my_var = 'Either use single quotes'
my_var = "or use double quotes"
# Boolean
my_var = True
my_var = False
```
Every variable can hold every type of data, so we don't need to specify the type at initialization.

We can have lists of several items:
```python
# List of numbers
my_list = [2.5, 7, 3.3]
# List of strings
my_list = ['E.T.', 'phone', 'home']
# List of booleans
my_list = [True, True, False]
```
Lists are 0-indexed and can be accessed with square brackets:
```python
# Set the 3rd element of my_list to 5
my_list[2] = 5
# Retrieve the 2nd element of my_list
new_var = my_list[1]
```
Similarly we can make tuples, which are just like lists, but their values cannot be modified after they have been initialized:
```python
# This is a tuple of three numbers
my_tuple = (3, 8.6, 4)
# Access to a tuple is similar to lists
new_var = my_tuple[1]
# But we cannot write to a tuple!
```
For numbers we have the following basic operators:
```python
# Define the two variables
a = 5
b = 2
# Addition
a + b
>> 7
# Subtraction
a - b
>> 3
# Multiplication
a * b
>> 10
# Division
a / b
>> 2.5
# Integer division
a // b
>> 2
# Modulo (remainder of integer division)
a % b
>> 1
# Power
a ** b
>> 25
```
We can have conditional statements
```python
if a < b:
    # printing is done with the print function
    print('a is smaller than b')
else:
    print('a is bigger than or equal to b')
```
<div class="alert alert-info">

**Note:** Indentation is important in Python. Writing a conditional statement, a loop, a function, etc., without an indent (`Tab`) raises an error.
</div>

with the following comparison operators:
```python
<  # Smaller than
<= # Smaller or equal
>  # Bigger than
>= # Bigger than or equal to
== # Equal to
!= # Not equal to
```
A `for` loop is declared as:
```python
# Note: range(n) creates numbers from 0 to n-1
for i in range(10):
    # Another for loop inside the for loop - we call this 'nested for loops'
    for j in range(10):
        # We can also have formatted strings indicated by a preceeding 'f', where the number i is inserted into the string
        print(f'Loop iteration number {j}')
```

Similarly a `while` loop:
```python
i = 0
while i < 10:
    print(f'Loop iteration number {i}')
    i = i + 1
```
And finally, a function can be defined using the `def` keyword:
```python
# A function that adds two numbers and returns the result
def addition(a, b):
    # Add the numbers
    result = a + b
    # Return the result
    return result

# Run the function
addition(5, 10)
>> 15
```

Experiment with Python in the cell below. Learn to do the simple things that come to mind and take the chance to explore all of the options above.

```python
# YOUR CODE HERE
```

### 2.2. NumPy
[Back to table of contents](#Table-of-contents)

[NumPy](https://NumPy.org/) library provides multidimensional matrix/array processing tools. Most image processing libraries in Python use NumPy arrays to represent images. 

It is best practice to import the NumPy library as `np`, to avoid typing `numpy` every time. This also makes the code more readable. Run the cell below to import NumPy.

```python
# Import the NumPy library
import numpy as np
```

To create a new NumPy array filled with zeros we can use `np.zeros`. It requires the size of the array we want to get as a tuple `(height, width)`. Run the cell below to see an example.

Optionally, we can specify the internal numerical [type](https://NumPy.org/devdocs/user/basics.types.html) of the array. For example, one can use 8-bit unsigned integers (`np.uint8`). This type is convenient for 8-bit images, as the resulting array will occupy less memory than with other types, such as `np.float32` or `np.float64`. For most of the labs, you will not need to mind the numerical type of your variables: it adapts automatically to represent what you are storing in each instance.

```python
# A NumPy array of size 5x10 with numerical type "unsigned 8-bit integer"
my_arr = np.zeros((5,10), dtype=np.uint8)
# Print the result
print(my_arr)
```

NumPy arrays can be accessed using the same notation used for lists, but with an index for every dimension. 
<div class="alert alert-warning">
    
**Beware:** In NumPy arrays, the first index specifies the row (y-axis location), not the x-axis location, while the second index specifies the column (x-axis location), not the y-axis location.
</div>

```python
# Set the value at the Yth row, Xth column to 1
my_arr[3, 7] = 1
print(my_arr)
```

We can also use *slicing* `:` to determine whole regions of the array to read or assign. 

```python
# Insert a rectangle with the upper left corner at (Y_1,X_1) and the lower right corner at (Y_2,X_2)
my_arr[:3, 2:6] = 1
print(my_arr)
```

As you can see in the code, in slicing, we define an interval of numbers by `<lower_limit>:<upper_limit + 1>`. If one of the limits is not specified, it is interpreted as _from the beginning_ or _to the end_. Note that the lower limit is inclusive, while the upper limit is exclusive.

Using slicing we can also extract patches of an array to create new arrays. Run the example below and see its effect.

```python
# Let's extract the rectangle that we created
my_rect = my_arr[:3, 2:6]
# Note: '\n' creates a new line when printing
print(f'Original array: \n{my_arr}\n\nExtracted patch: \n{my_rect}')
```

Using NumPy arrays, we can easily perform the basic operations element-wise on all elements of one or several arrays (). Run the cell below to define a random array to work with.

```python
# First we define an array of size 5x5 filled with random integers from 0 to 10
uid = 42  # Define the seed value
np.random.seed(uid)
rand_array = np.random.randint(low=0, high=10, size=(5,5))
print(f'Initial array:\n{rand_array}')
```

We can perform any operation with a scalar for all elements of the array at the same time by using the array as if it was a variable with a single number. Run the example below and make sure you understand its results.

```python
# Add 2 to all elements
rand_array = rand_array + 2
print(f'rand_array + 2:\n{rand_array}\n')
# Divide all elements of the resulting array by 2
rand_array = rand_array / 2
print(f'rand_array / 2:\n{rand_array}')
```

Similarly, we can use expressions with multiple arrays of the same size and any of the basic operators: the calculation will be made on the co-positioned elements of the different arrays. Run the example below and make sure you understand its results.

```python
# Let's initialize two new random arrays
rand_array_1 = np.random.randint(low=0, high=10, size=(5,5))
rand_array_2 = np.random.randint(low=0, high=10, size=(5,5))
print(f'Array 1:\n{rand_array_1}\n\nArray 2:\n{rand_array_2}')
# We can for example add them together
added_arrays = rand_array_1 + rand_array_2
print(f'\nArray 1 + Array 2:\n{added_arrays}')
# Or multiply them
multiplied_arrays = rand_array_1 * rand_array_2
print(f'\nArray 1 * Array 2:\n{multiplied_arrays}')
```

We can also perform _advanced indexing_ of NumPy arrays: we can provide a [boolean](https://en.wikipedia.org/wiki/Boolean_data_type) array as the index and extract/modify only the values that are specified as `True`. An application of this form of indexing could be to take the multiplied array from the previous cell and set all its elements that are larger than $25$ to $100$. Run the example below and make sure you understand its results.

```python
# First print the multiplied array as a reference
print(f'Multiplied array:\n{multiplied_arrays}')
# Create the boolean array that contains true at every location, where the multiplied array has a value bigger than 25
boolean_array = multiplied_arrays > 25
print(f'\nBoolean array:\t(Multiplied array > 25)\n{boolean_array}')
# Now we can use this array to index the multiplied array and set all elements indicated by True to 100
multiplied_arrays[boolean_array] = 100
print(f'\nMultiplied array with all elements > 25 set to 100:\n{multiplied_arrays}')
```

Note that you can also create the boolean array in place, without explicitly declaring an additional variable. Run the cell below and see the difference with the preceeding cell. We will reuse the variable `multiplied_arrays`, and set all elements smaller than $20$ to $1$.

```python
# We get the array we have, and index only the elements smaller than 20.
multiplied_arrays[multiplied_arrays < 20] = 1
print(f'\nMultiplied array with all elements < 20 set to 1:\n{multiplied_arrays}')
```

### 2.3. Matplotlib
[Back to table of contents](#Table-of-contents)

[Matplotlib](https://matplotlib.org/), specifically, [PyPlot module](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.html#module-matplotlib.pyplot) allows us to load images as NumPy arrays, and visualize NumPy arrays as images. 
 

It is best practice to import the module as `plt`.

```python
# Set the image display to be interactive 
%matplotlib widget
# Import pyplot
import matplotlib.pyplot as plt
```

The only functionality you really need to know is how to load images using `plt.imread`:

```python
# Load an example image located at 'images/boats.tif'
example_img = plt.imread('images/boats.tif')
# Then, we can check the image is really a NumPy array
type(example_img)
```

To see the image, use `plt.imshow`.

```python
# When using plt.imshow we first need to declare a new figure
plt.figure()
# Then we add the image to the figure
plt.imshow(example_img, cmap='gray')
# Give it a title
plt.title(f"Example Image - BOATS")
# And finally we make sure it is displayed on the screen
plt.show()
```

Experiment with Python, NumPy and Matplotlib in the empty cell below. Learn to do the simple things that come to mind and take the chance to explore all of the options above.

```python
# Example code to create image patch of the boat
boat_patch = example_img[150:525, 200:476]
plt.figure()
plt.imshow(boat_patch, cmap='gray')
plt.title(f"Patch image of the boat")
plt.show()
```

Loading multiple images

```python
# Configure plotting as dynamic
%matplotlib widget
# Import required packages for this exercise
import numpy as np
import matplotlib.pyplot as plt

# Import images
car   = plt.imread('images/car_pad.tif')
hrct  = plt.imread('images/hrct.tif')
hands = plt.imread('images/hands.tif')
plate = plt.imread('images/plate.tif')
boats = plt.imread('images/boats.tif')
```

View the imported images in a grid format.

```python
# Create a grid layout for displaying images
fig, axs = plt.subplots(2, 2, figsize=(8, 8))  # 2 rows, 2 columns
fig.suptitle("Image Grid", fontsize=16)  # Add a title for the grid

# List of images and titles
images = [car, hrct, hands, plate, boats]
titles = ['Car', 'HRCT', 'Hands', 'Plate']

# Loop to display each image
for i, ax in enumerate(axs.flat):
    if i < len(images):  # Ensure we don't run out of images
        ax.imshow(images[i], cmap='gray')  # Display image
        ax.set_title(titles[i])  # Add a title
    ax.axis('off')  # Hide axes for clarity

plt.show()
```

### Practice question

Answer the following question.

In the image `car`, what is the mean value of **the carpet that the car is standing on**? 
Hint: select multiple patches from the carpet area. Ensure the patches are located away from the car and the edges of the image. Use np.mean to compute the average value of each patch and then calculate the overall mean.

1. Around 40
2. Around 60 
3. Around 80

<!-- grade_id: cell-40970ade3d3b9d3e_2022, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE


```

# 3. Image Thresholding
[Back to table of contents](#Table-of-contents)

In the next cells, you will see a very basic example. We will apply a pixel-wise operation on the image `car`: all the pixels with a value below a treshold (given in $\%$) will be set to the *maximum* value. 

First, we define the thresholding function `your_function`. Run the next cell to define it.

```python
# Define your function
def your_function(image, threshold):    
    # We make a copy of the original, where we will apply our threshold.
    output = np.copy(image)
    # For greater flexibility, we will get the value as a percentage of the maximum
    value = threshold * 0.01 * np.amax(image)
    # Apply threshold
    output[image < value] = np.amax(image)
    return output
```

```python
# You can play around with the threshold value
threshold = 60
```

```python
# Thresholding the car image
thresholded_image = your_function(car,threshold)
plt.figure()
plt.imshow(thresholded_image, cmap='gray')
plt.title(f"Thresholded image of the car")
plt.show()
```

## 3.1 Range thresholding

Now you will program a function to apply range thresholding to the image.

Your first task is to define a function `range_function(img, rmin, rmax)` that will only **keep** pixel values inside a range of values $[r_{\mathrm{min}}, r_{\mathrm{max}}]$ ($r_{\mathrm{min}}$ and $r_{\mathrm{max}}$ included), while the values outside this range will be set to $0$.
You will apply it on the image `car`, so `rmin` and `rmax` should be between $0$ and $255$. Define this function in the following cell. 

<div class="alert alert-info">
  
<b>Note: </b> Use boolean indexing of the NumPy array `img`, as in the thresholding example, to perform the necessary operations.
</div>

<!-- grade_id: cell-f7d3dd947484c706_2022, points: 0, grade: false, solution: true -->
```python
# Define your function
def range_function(img, rmin, rmax):    
    # Best practices are creating a copy. You don't wanna mess the original.
    output = np.copy(img)
    # set the pixels of output that are smaller than rmin or bigger than rmax to 0
    # YOUR CODE HERE
    return output
```

You can test your function with the following cell. In it, we define a linear array `test_array` that is formed by the integers from 0 to 255 in ascending order, using the method [NumPy.arange](https://NumPy.org/doc/stable/reference/generated/NumPy.arange.html). Then, we apply your function with the range $[100, 199]$ to it. There should be exactly $100$ elements that were not set to $0$. If the cell runs smoothly, your implementation is probably correct.
<div class="alert alert-info">
    
**Note:** The function `np.count_nonzero(arr)` counts the number of elements in `arr` that are not zero.
</div>

<!-- grade_id: cell-d32bcff0667c88a9_2022, points: 1, grade: true, solution: false -->
```python
test_array = np.arange(0, 256)
assert np.count_nonzero(range_function(test_array, 100, 199)) == 100, 'Your implementation is not yet correct.'
print('Congratulations! Your function is correct.')

```

```python
# test the range thresholding on the car image
# YOUR CODE HERE
```

# 4. Image processing in Python using scikit-image
[Back to table of contents](#Table-of-contents)

Throughout this course we will focus on one most common image processing library
[**`scikit-image`**](https://scikit-image.org/).

Scikit stands for SciPy Toolkits, it is developed specifically for image-processing tasks. These toolkits are add-on packages for SciPy. They have become very well known and have a large community of users and developers. 

You can find the documentation for all the API's available over [here](https://scikit-image.org/docs/stable/api/api.html)

Just as with `NumPy` (`np`) or `matplotlib.pyplot` (`plt`), there is a standard abbreviation to import this library. 

```python
# Scikit-image
import skimage
```

# 4.1. Examples
[Back to table of contents](#Table-of-contents)

Now we will go through two very basic operations: rotation and edge detection for you to get familiar with scikit-image. We will show you the basic syntax. For the full list of parameters and options, go to the official documentation of each function. 


### Rotation in Scikit Image

The basic syntax of the function for grayscale images is  (see [the documentation](https://scikit-image.org/docs/stable/api/skimage.transform.html?highlight=rotate#skimage.transform.rotate)):
```python
output = skimage.transform.rotate(image, angle, resize=False, center=None, mode='reflect', preserve_range=True)
```

The parameters are:
* `input`: The source image, 
* `angle`: Rotation angle in degrees,
* `resize`: Specify whether to resize the image to contain the full original. Defaults to `False`,
* `center`: Specifies the center of rotation. The order is (columns, rows),
* `mode`: Specify boundary conditions. 


We will apply sciki-image's rotation to `car` for an angle of $60^\circ$ with the center of rotation at $(50, 50)$. Run the next cell to see the result. Modify the cell and experiment with the different parameters.

```python
# First, we import the module transform
from skimage import transform

# Now we apply operation
rot_car_skimage = skimage.transform.rotate(car, angle=60, resize=False, center=(50, 50), mode='reflect', preserve_range=True)

# Visualize the original and transformed images
plt.close('all')
plt.figure(figsize=(8, 5))

# Original image
plt.subplot(1, 2, 1)
plt.imshow(car, cmap='gray')
plt.title('Original Image')
plt.axis('off')

# Blurred image
plt.subplot(1, 2, 2)
plt.imshow(rot_car_skimage, cmap='gray')
plt.title('Transformed Image')
plt.axis('off')

plt.tight_layout()
plt.show()

```

### Edge Detection in Scikit Image

The basic syntax of the Sobel filter for grayscale images is  (see [the documentation](https://scikit-image.org/docs/dev/api/skimage.filters.html#skimage.filters.sobel)):

```python
output = skimage.filters.sobel(image, axis=None, mode='reflect')
```

The parameters are:
* `image`: The source image, 
* `axis`: If given, compute the axis in the specified direction. If `None`, the magnitude of the gradient is calculated,
* `mode`: Specify boundary conditions (see example above). 


Run the next cell to apply Scikit-Image Sobel filter on the image `plate`. Modify the cell and play with the different values and with different images.

```python
# First, we import the module filters
from skimage import filters

# Apply filter
plate_edges_skimage = skimage.filters.sobel(plate, axis=None, mode='reflect')

# Visualize the original and filtered images
plt.close('all')
plt.figure(figsize=(8, 5))

# Original image
plt.subplot(1, 2, 1)
plt.imshow(plate, cmap='gray')
plt.title('Original Image')
plt.axis('off')

# Blurred image
plt.subplot(1, 2, 2)
plt.imshow(plate_edges_skimage, cmap='gray')
plt.title('Filtered Image')
plt.axis('off')

plt.tight_layout()
plt.show()
```

## 4.2. Try it yourself! 
[Back to table of contents](#Table-of-contents)

Complete the function `gaussian_blur`, which implements a **Gaussian smoothing** on the given input image. Set
* the standard deviation of the Gaussian kernel to be $5$,
* boundary conditions to be reflective, and
* keep the original range of values.

<div class='alert alert-warning'>
    <b>Note:</b> In python it's important to always specify the parameter names when using a function. Only this way, you can be sure that you set the correct parameters. For example, don't use <code>skimage.filters.sobel(plate, None, 'reflect')</code>, but instead use <code>skimage.filters.sobel(plate, axis=None, mode='reflect')</code>.
</div>

<!-- grade_id: cell-3aad12f740231a32_2022, points: 0, grade: false, solution: true -->
```python
# Function that applies a Gaussian blur
def gaussian_blur(img):
    out = np.zeros(img.shape)
    
    # YOUR CODE HERE

    return out
```

<!-- grade_id: cell-1919408b3ba0b1c7_2022, points: 2, grade: true, solution: false -->
```python
# Visualize the original and filtered images
plt.close('all')
plt.figure(figsize=(8, 5))

# Original image
plt.subplot(1, 2, 1)
plt.imshow(boats, cmap='gray')
plt.title('Original Image')
plt.axis('off')

# Blurred image
plt.subplot(1, 2, 2)
plt.imshow(gaussian_blur(boats), cmap='gray')
plt.title('Filtered Image')
plt.axis('off')

plt.tight_layout()
plt.show()
```

<div class="alert alert-success">

Congratulations on finishing this first lab of ECE172A!
</div>

```python

```